--- 
courseTitle       : Insert the course title here
chapterTitle      : Insert the chapter title here
description       : What is this chapter about? Add here your description
framework : datamind
mode: selfcontained
--- 

# START READ-AND-DELETE 
  1. Complete YAML front matter above
  2. Write using R Markdown (http://en.wikipedia.org/wiki/Markdown)
  3. Use an empty line followed by three dashes to separate exercises!
  4. Consult the documentation at: http://www.datamind.org/#/CC/help if you're stuck
# END READ-AND-DELETE

## The title of the first exercise

In the assignement, you can explain statistical concepts or R functionality. Assignments are written in (R) Markdown.

You can emphasize important words by putting them in **bold** or *italics*, or indicate that some parts are `r_code` with backticks. You can even add lists, to structure your explanation: 
- item one
- item two

Or you could link to interesting content: [this is a link](http://en.wikipedia.org/wiki/Markdown), or [another useful link](http://www.rstudio.com/ide/docs/authoring/using_markdown...).

*** =instructions
  - Provide one instruction to your students/employees/... on what specific problem they should solve in this exercise
  - Provide a second instruction here.
  - ect.

*** =hint

Hints help the students to find a correct solution for the exercise. Hints are optional, but their usage is highly recommended.

*** =sample_code

```{r}
# The sample code is what is shown in the students editor by default
```

*** =solution

```{r}
# The solution code solves the assignement. Make sure the solution passes the submission correctness test (we will check this;-)!
```

*** =sct

```{r}
# The Submission Correctness Test code
# It takes as input the students workspace, his code and what was printed to the console during the exercise.
# Based on these inputs the SCT returns:
# (i) whether the students code was correct
# (ii) a message to the student to either help him find what was wrong or congragulate him on finding the solution.
```

*** =pre_exercise_code

```{r}
# Pre-exercise code is executed before the exercise starts (what's in a name...)
# It can be used to:
# Pre-load a dataset
# Render graphs that are shown at the start of the exercise
# ...
```

---

## Title of the second exercise!

Assigment here.

*** =instructions

- Instructions here
  
*** =hint

Hint here. 

*** =sample_code

```{r}
# Sample code here
```

*** =solution

```{r}
# Solution code here 
```

*** =sct
```{r}
# Submission Correctness Test here
```

*** =pre_exercise_code
```{r}
# Pre-exercise code here
```