--- 
courseTitle       : Insert the course title here
chapterTitle      : Insert the chapter title here
description       : What is this chapter about? Add here your description
framework : datamind
mode: selfcontained
--- 

# START READ AND DELETE 
  1. Edit YAML front matter above
  2. Write using R Markdown (http://en.wikipedia.org/wiki/Markdown)
  3. Use an empty line followed by three dashes to separate exercises!
  4. Consult the documentation at: http://www.datamind.org/#/CC/help if you're stuck
# END READ AND DELETE

## The title of the exercise

In the assignement, you can explain statistical concepts or R functionality. Assignments are written in (R) Markdown.

You can emphasize important words by putting them in **bold** or *italics*, or indicate that some parts are `r_code` with backticks. You can even add lists, to structure your explanation: 
- item one
- item two

Or you could link to interesting content: [this is a link](http://en.wikipedia.org/wiki/Markdown), or [another useful link](http://www.rstudio.com/ide/docs/authoring/using_markdown...).

*** =instructions
  - Provide one instruction to your students/employees/... on what specific problem they should solve in this exercise
  - Provide a second instruction here.
  - ect.

*** =hint

Hints help the students to find a correct solution for the exercise. Hints are optional, but their usage is highly recommended.

*** =sample_code

```{r}
# The sample code is what is shown in the students editor by default
```

*** =solution

```{r}
# The sample code is what is shown in the students editor by default
```

*** =sct

```{r}
DM.result <- DM.codeTest("matrix(1:9, byrow=TRUE, nrow=3)");
```

---

## A second exercise starts here everything below is analogous to what's explained above.

Assigment here.

*** =instructions

- Instructions here
  
*** =hint

Hint here. 

*** =sample_code

```{r}
# Sample code here
```

*** =solution

```{r}
# Solution code here 
```

## sct
```{r}
# Submission Correctness Test here
```